// primitive_types2.rs
// 补充不完整代码行的缺失部分！
// 没有提示，也没有什么诀窍，只要习惯于键入这些内容就可以了 :)

// I AM NOT DONE

fn main() {
    // Characters (`char`)

    let my_first_initial = 'C';
    if my_first_initial.is_alphabetic() {
        println!("Alphabetical!");// 译："字母！"
    } else if my_first_initial.is_numeric() {
        println!("Numerical!");// 译："数字！"
    } else {
        println!("Neither alphabetic nor numeric!");// 译："既不是字母也不是数字！"
    }

    let // 像上面的示例一样完成这行代码！你最喜欢的角色是什么？
    // 试试一个字母，或者一个数字，也可以一个特殊字符，又或者一个不属于
    // 你母语的字符，一个表情符号看起来也不错。
    if your_character.is_alphabetic() {
        println!("Alphabetical!");
    } else if your_character.is_numeric() {
        println!("Numerical!");
    } else {
        println!("Neither alphabetic nor numeric!");
    }
}
