// primitive_types1.rs
// 补充不完整代码行的缺失部分！
// 没有提示，也没有什么诀窍，只要习惯于键入这些内容就可以了 :)

// I AM NOT DONE

fn main() {
    // Booleans (`bool`)

    let is_morning = true;
    if is_morning {
        println!("Good morning!");// 译："早上好"
    }

    let // 参照上面的示例来补充这一行的缺失部分！或者让它值为 false ！
    if is_evening {
        println!("Good evening!");// 译："晚上好"
    }
}
