// variables1.rs
// 让我能够编译！执行 `rustex hint variables1` 获取提示 :)

// 关于 `I AM NOT DONE`：
// 即使你已经想到了解决方案，我们有时也会鼓励你在特定的练习
// 继续去尝试做到更好。如果你已经通过练习并做好了应对下一题
// 的准备, 请删除下面的 `I AM NOT DONE` 注释。

// I AM NOT DONE

/// 翻译: [mg-chao](https://github.com/mg-chao)
/// 加油 💪
fn main() {
    x = 5;
    println!("x has the value {}", x);
}
