// variables5.rs
// 让我能够编译！执行 `rustex hint variables5` 获取提示 :)

// I AM NOT DONE

/// 翻译: [mg-chao](https://github.com/mg-chao)
fn main() {
    let number = "T-H-R-E-E"; // don't change this line
    println!("Spell a Number : {}", number);// 译："拼接的数字：{}"
    number = 3;
    println!("Number plus two is : {}", number + 2);// 译："数字加上二是：{}"
}
