// primitive_types3.rs
// 在 ??? 处创建一个不少于 100 个元素的数组。
// 执行 `rustex hint primitive_types3` 获取提示！

// I AM NOT DONE

fn main() {
    let a = ???

    if a.len() >= 100 {
        println!("Wow, that's a big array!");// 译："哇！那数组可真大！"
    } else {
        println!("Meh, I eat arrays like that for breakfast.");// 译："嗯，我把这样的数组当早餐吃。"
    }
}
